---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

Lesson 2: Introduction to Essay Writing

This is the second course in Academic English: Writing. By introducing you to three types of academic essays, this course will specifically prepare you for university study, but anyone who wants to improve their writing skills can benefit from this course.

After completing this course, you will be able to:
-Create a valid thesis statement for your thesis
-Plan and write comparisons / contrasts, cause and effect / arguments
-Write perfect text paragraphs

Note: Lectures and practice activities are free, but you must upgrade to the paid version to take the quiz and get feedback on writing assignments.
syllabus
Course Introduction
-This is academic English: the second course in writing [paper on technology](https://www.wowessays.com/topics/technology/). In the previous lesson, you reviewed sentence types and punctuation. In this course, you will use this information to improve your writing skills. In this course, you will learn everything you need to know about academic paper writing, especially how to write three types of papers: comparison / contrast, cause / effect / argument, and argument. To pass this course, you need to pass all four quizzes and pass all three writing assignments. After completing one activity, you can move on to the next. Please enjoy!

Prose Writing
-In this module, you will begin to study the structure of the article and other important tools to improve the quality of writing. There is a lot of information in this module, but it is necessary to write well. Make sure you take notes so that you remember these tools when writing your dissertation. Learner's Note: This course is designed for English learners with intermediate English writing skills. The sample papers in this course are for this level. However, the principles discussed in this course are practical for writers of any level. If your level is lower, try your best. If you are a more advanced writer, feel free to write more complex and complicated articles than the examples. Just make sure you follow the structure of the introduction.

Writing a Comparison / Contrast Paper
-Now you can write your first academic paper-Comparison / Contrast paper. In this module, you will learn what type of thesis and how to write it. You will then look at some examples and practice writing your own comparison / contrast articles. Remember that the sample articles in this lesson are typical examples of intermediate students. Write a comparison / contrast article that suits your writing abilities. good luck!

Writing a Cause and Effect Paper
-Now you will learn to write a causal dissertation. This is another academic paper that you may be asked to write in a college class. For this type of article, you will consider why something happened or the impact of something. The sample articles in this module also represent intermediate writers. Write a causal thesis suitable for your own level of English. Just remember to follow the advice given in the course.

Writing argument
-In the last module, you will learn how to write the most common types of college thesis. Argument papers are also probably the most interesting ones. In this article, you will try to convince readers of your argument or position on a controversial topic. You must consider good reasons to support your position. Remember to write an argument for what you can.

`@instructions`
Becoming a better writer can never go wrong. Whether you are preparing for college, drafting a resume and cover letter, writing a sales email, or trying to retain your memory, you need to be able to communicate in writing. Fortunately, free lessons abound and can guide you in writing.

`@hint`


`@pre_exercise_code`
```{r}
#include <stdio.h>
int main() {    
    int number1, number2, sum;
    
    printf("Enter two integers: ");
    scanf("%d %d", &number1, &number2);
    // calculating sum
    sum = number1 + number2;      
    
    printf("%d + %d = %d", number1, number2, sum);
    return 0;
```

`@sample_code`
```{r}
#include <stdio.h>
int main() {    
    int number1, number2, sum;
    
    printf("Enter two integers: ");
    scanf("%d %d", &number1, &number2);
    // calculating sum
    sum = number1 + number2;      
    
    printf("%d + %d = %d", number1, number2, sum);
    return 0;
```

`@solution`
```{r}

```

`@sct`
```{r}

```
